```python
# Initialize Otter
!pip install -q otter-grader &> /dev/null
import otter
grader = otter.Notebook("python_quiz.ipynb")
```

In order for us to:

1. get a sense of how strong your python is
1. give you an example of how many of the assignments will work

Please give it an honest attempt. 

## **Question 1.** Extract Student Records from JSON

Read the provided JSON file "records.json" and extract a list of student records, with each record being a dictionary.

```python
import json

with open('records.json', 'r') as file:
    data = json.load(file)
    records = data['students']

print(records)

```

Output:
```
[{'name': '     Alice Johnson ', 'age': 20, 'courses': ['CS501', 'CS502', 'CS503']}, {'name': 'Bob    Smith ', 'age': 22, 'courses': ['CS505', 'CS504']}, {'name': '    Charlie Davis', 'age': 19, 'courses': ['CS501', 'CS508', 'CS506']}, {'name': '   Diana    Taylor', 'age': 21, 'courses': ['CS503', 'CS502', 'CS506', 'CS507']}, {'name': '   Bob   Dylan ', 'age': 21, 'courses': ['CS509', 'CS502', 'CS504']}]
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

## **Question 2.** Copy JSON Content Using shutil

Create a new JSON file called `clone-records.json` and copy the contents of 'records.json' into it using the shutil library.

```python
import shutil

source_file = 'records.json'
destination_file = 'cloned-records.json'

shutil.copyfile(source_file, destination_file)
```

Output:
```
'cloned-records.json'
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

## **Question 3.** Sort Students by Age in Descending Order

Read the provided JSON file "records.json" and extract a list of student records and assign it to `sorted_records`, with each record being a dictionary. Use only one line of python.

```python
sorted_records = sorted(records, key=lambda x: x['age'], reverse=True)

print(sorted_records)

```

Output:
```
[{'name': 'Bob    Smith ', 'age': 22, 'courses': ['CS505', 'CS504']}, {'name': '   Diana    Taylor', 'age': 21, 'courses': ['CS503', 'CS502', 'CS506', 'CS507']}, {'name': '   Bob   Dylan ', 'age': 21, 'courses': ['CS509', 'CS502', 'CS504']}, {'name': '     Alice Johnson ', 'age': 20, 'courses': ['CS501', 'CS502', 'CS503']}, {'name': '    Charlie Davis', 'age': 19, 'courses': ['CS501', 'CS508', 'CS506']}]
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

## **Question 4.** Clean and Format Student Names Using Regex

Write Python code to format the 'name' field in each dictionary from the list to the 'Firstname Lastname' format. Remove any extra white spaces between the first name and last name, as well as leading and trailing spaces in each name string. Store the result in `cleaned_records`. (Hint: Use Regex)

```python
import re


print(cleaned_records)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[8], line 4
      1 import re
----> 4 print(cleaned_records)

NameError: name 'cleaned_records' is not defined
```

```python
grader.check("q4")
```

Output:
```
q4 results:
    q4 - 1 result:
        ❌ Test case failed
        Trying:
            len(cleaned_records) == 5 and type(cleaned_records) == list
        Expecting:
            True
        **********************************************************************
        Line 1, in q4 0
        Failed example:
            len(cleaned_records) == 5 and type(cleaned_records) == list
        Exception raised:
            Traceback (most recent call last):
              File "/Users/williamli/anaconda3/lib/python3.11/doctest.py", line 1351, in __run
                exec(compile(example.source, filename, "single",
              File "<doctest q4 0[0]>", line 1, in <module>
                len(cleaned_records) == 5 and type(cleaned_records) == list
                    ^^^^^^^^^^^^^^^
            NameError: name 'cleaned_records' is not defined
```

## **Question 5.** Add Course Count to Each Student Record

Create a new dataset called `courses_records` where each student record includes an attribute called `num_courses` representing the total number of courses the student is enrolled in. 

```python

courses_records = [{**record, 'num_courses': len(record['courses'])} for record in records]
print(courses_records)

```

Output:
```
[{'name': '     Alice Johnson ', 'age': 20, 'courses': ['CS501', 'CS502', 'CS503'], 'num_courses': 3}, {'name': 'Bob    Smith ', 'age': 22, 'courses': ['CS505', 'CS504'], 'num_courses': 2}, {'name': '    Charlie Davis', 'age': 19, 'courses': ['CS501', 'CS508', 'CS506'], 'num_courses': 3}, {'name': '   Diana    Taylor', 'age': 21, 'courses': ['CS503', 'CS502', 'CS506', 'CS507'], 'num_courses': 4}, {'name': '   Bob   Dylan ', 'age': 21, 'courses': ['CS509', 'CS502', 'CS504'], 'num_courses': 3}]
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

## **Question 6.** Filter Students by Course Enrollment

Create a function that accepts a course name as input and returns a list of students who are enrolled in that course.

```python
def get_students(course_name):

    return [student['name'] for student in records if course_name in student['courses']]

```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

## **Question 7.** Calculate Average Age of Students

Calculate and return the average age of the students, rounding off to the nearest integer and assign it to a variable called `avg_age`.

```python
import numpy as np

def get_avg_age(student_list):
    ages = [student['age'] for student in student_list]
    return int(np.round(np.mean(ages)))

avg_age = get_avg_age(records)
print(avg_age)

```

Output:
```
21
```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

## **Question 8.** List of Unique Courses

Create a list of unique courses offered, as per the student records, and assign it to `total_courses`.

```python
import numpy as np

def get_unq_courses(student_list):
    courses = set()
    for student in student_list:
        courses.update(student['courses'])
    return np.array(list(courses))

total_courses = get_unq_courses(records)
print(total_courses)

```

Output:
```
['CS501' 'CS504' 'CS503' 'CS507' 'CS506' 'CS502' 'CS505' 'CS508' 'CS509']
```

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
```

## **Question 9.** Group Students by Age 

Return a dictionary where each key is an age, and the value is a list of student names who are of that age, and assign it to a variable called `grouped_students`. Utilize itertools.groupby to group the list of students by their age.

```python
import itertools

def group_students_by_age(student_list):

    grouped_students = {}
    for student in student_list:
        age = student['age']

print(grouped_students) 
```

```python
grader.check("q9")
```

## **Question 10.** Generate Course Enrollment Summary

Generate a dictionary where each key is a course name and the value is the number of students enrolled in that course.


```python
from collections import Counter

def get_course_dict(student_records):

course_enrollment = ...
print(course_enrollment)
```

```python
grader.check("q10")
```

Tests for q1 (points: 10)
```
>>> type(records) == list
True
```


Tests for q10 (points: 10)
```
>>> type(course_enrollment) == dict
True
```


Tests for q2 (points: 10)
```
>>> import os
>>> os.path.exists('cloned-records.json') and os.path.getsize('cloned-records.json') > 0
True
```


Tests for q3 (points: [0])
```
>>> len(sorted_records) == 5
True
```


Tests for q4 (points: 10)
```
>>> len(cleaned_records) == 5 and type(cleaned_records) == list
True
```


Tests for q5 (points: 10)
```
>>> len(courses_records) == 5 and type(courses_records) == list
True
```


Tests for q6 (points: 10)
```
>>> len(get_students('CS506')) == 2
True
```


Tests for q7 (points: 10)
```
>>> type(avg_age) == int
True
```


Tests for q8 (points: 10)
```
>>> type(total_courses) == np.ndarray
True
```


Tests for q9 (points: 10)
```
>>> type(grouped_students) == dict
True
```
